Bitte gehen Sie sicher, dass das Notebook komplett durchläuft bevor Sie die Lösung einreichen!.
Wählen Sie dazu **Kernel -> Restart & Run All**

Füllen Sie alle Stellen aus, an denen Sie `Your Code HERE` oder "YOUR ANSWER HERE" finden. Schließlich tragen Sie Ihren Namen ein in der nächsten Zelle ein.

```python
NAME = "Florian Krax"
```

---

# 1) Arrays erzeugen

Zum Aufwärmen gehen wir ein paar Möglichkeiten durch Numpy-Arrays zu erzeugen

Dafür muss jetzt natürlich erst Numpy importiert werden:

```python
import numpy as np
```

## 1a)

Python Listen können einfach in Numpy Arrays umgewandelt werden. Wie?

<!-- grade_id: cell-array, points: 0, grade: false, solution: true -->
```python
array = [1,2,3,4]
# YOUR CODE HERE
array = np.asarray(array)
```

<!-- grade_id: cell-array-test, points: 1, grade: true, solution: false -->
```python
assert isinstance(array, np.ndarray) # Testen ob Array nun ein Numpy Array ist
assert array.shape == (4,)           # Testen ob der shape des neuen Array passt
```

## 1b)

Erzeuge Arrays mit folgenden Eigenschaften:
- `intrange` eine Numpy Variante von `range(5, 10, 2)`
- `numrange` ein Numpy Array mit 101 **gleichverteilten** Werten zwischen **2** und **4**

<!-- grade_id: cell-ranges, points: 0, grade: false, solution: true -->
```python
intrange = None
numrange = None
# YOUR CODE HERE
intrange = np.arange(5,10,2)
numrange = np.linspace(2.0, 4.0, num=101)
```

<!-- grade_id: cell-ranges-tests, points: 1, grade: true, solution: false -->
```python
assert isinstance(intrange, np.ndarray)
assert isinstance(numrange, np.ndarray)
assert intrange.shape == (3,)
assert numrange.shape == (101,)
assert all([result == wanted for result, wanted in zip(intrange, range(5, 10, 2))])
assert all([round(result, 6) == round(wanted / 100, 6)
            for result, wanted in zip(numrange, range(200, 401, 2))])
```

## 1c)

Manchmal benötigt man Arrays mit einheitlichen Werten erzeuge folgende:
- `zarray` eine **6x4** Matrix aus **Nullen**
- `oarray` eine **1x100** Matrix aus **Einsen**
- `varray` eine **4x4** Matrix gefüllt mit **4**
- `darray` eine **4x4** Matrix aus Nullen mit den Diagonalelementen `[1,2,3,4]`

<!-- grade_id: cell-znfdarray, points: 0, grade: false, solution: true -->
```python
zarray = None
oarray = None
varray = None
darray = None
# YOUR CODE HERE
zarray = np.zeros((6,4))
oarray = np.ones((1,100))
varray = np.full((4,4),4)
darray = np.identity(4) * np.arange(1,5)

```

<!-- grade_id: cell-znfdarray-test, points: 1, grade: true, solution: false -->
```python
assert isinstance(zarray, np.ndarray)
assert isinstance(oarray, np.ndarray)
assert isinstance(varray, np.ndarray)
assert isinstance(darray, np.ndarray)
assert zarray.shape == (6, 4)
assert oarray.shape == (1, 100)
assert varray.shape == (4, 4)
assert darray.shape == (4, 4)
```

# 2) Mathematische Operationen

Mit Numpy Arrays lassen sich einfach mathemische Operationen auf den kompletten Array anwenden. Schreibe eine Funktion welche folgende Operationen ausführt:

$f(x) = 5 + 3x + x^3 + \sqrt{x} + \mathrm{e}^x + \frac{1}{x} + \sin(2x)$

<!-- grade_id: cell-mathfuncs, points: 0, grade: false, solution: true -->
```python
x_points = None
# YOUR CODE HERE
intrange = np.arange(1,3,1)
def function(x: np.ndarray) -> np.ndarray:
    r"""
    """
    # YOUR CODE HERE
    val = 5 + 3 * x + np.power(x, 3) + np.power(x, 0.5) + np.exp(x) + 1 / x + np.sin(2 * x)
    print(val)
    return val
```

<!-- grade_id: cell-mathfuncs-test, points: 1, grade: true, solution: false -->
```python
assert round(function(1), 6) == 14.627579
assert round(function(2), 6) == 27.546467
```

Output:
```
14.627579255284726
27.546467165995818
```

# 3) Summen und Differenzen

Mit dem folgenden Array `[1, 2, 3, 4, 5, 6, 7, 8]`:
- `suma` ist die Summe aller Zahlen
- `proda` ist das Produkt aller Zahlen
- `csum` ist die Summe aller Zahlen bis zu diesem Eintrag, also `[1, 3, 6, 10, 15, 21, 28, 36]`
- `diffa` ist die Differenz benachbarter Zahlen

<!-- grade_id: cell-sumdiff, points: 0, grade: false, solution: true -->
```python
array = [1, 2, 3, 4, 5, 6, 7, 8]
suma = None
proda = None
csum = None
diffa = None
# YOUR CODE HERE
suma = np.sum(array)
proda = np.prod(array)
csum = np.cumsum(array)
diffa = np.diff(array)
```

<!-- grade_id: cell-sumdiff-test, points: 1, grade: true, solution: false -->
```python
assert suma == 36
assert proda == 40320
assert np.array_equal(csum, [1, 3, 6, 10, 15, 21, 28, 36])
assert np.array_equal(diffa, [1, 1, 1, 1, 1, 1, 1])
```

# 4) Ableitungen

Mit Numpy können Ableitungen von berechnet werden. Berechne die Ableitung einer Funktion $f(x)$ für
- $x$ = `[1, 2, 3, 4, 5, 6]`
- $f(x)$ = `[1, 4, 9, 16, 25, 36]`
- `grad` ist die gesuchte Ableitung

<!-- grade_id: cell-gradient, points: 0, grade: false, solution: true -->
```python
x = np.array([1, 2, 3, 4, 5, 6])
f_of_x = np.array([1, 4, 9, 16, 25, 36])
grad = np.gradient(f_of_x)
# YOUR CODE HERE
```

<!-- grade_id: cell-gradient-test, points: 1, grade: true, solution: false -->
```python
assert np.array_equal(grad, [3, 4, 6, 8, 10, 11])
```

# 5) Vektoren

In Numpy sind ebenfall Vektoroperationen enthalten. Die zwei häufigsten sind hier das Skalar- und das Kreuzprodukt. Wende folgende Operationen auf die Vektoren $v_1$ und $v_2$ an:
- `scalar` $v_1 \cdot v_2$
- `cross` $v_1 \times v_2$

<!-- grade_id: cell-vectors, points: 0, grade: false, solution: true -->
```python
v_1 = [1, 2, 3]
v_2 = [4, 2, 5]
scalar = np.dot(v_1, v_2)
cross = np.cross(v_1, v_2)
# YOUR CODE HERE
```

<!-- grade_id: cell-vectors-test, points: 1, grade: true, solution: false -->
```python
assert scalar == 23
assert np.array_equal(cross, [4, 7, -6])
```

```python

```